<?xml version="1.0"?>
<section xml:id="selection_the-in-and-not-in-operators">
  <title>Sequence Conditional Operators</title>
  <p>Along with basic comparsion operators, Python also provides built-in operators to create boolean expressions with sequences, such as strings and lists.</p>
  <p>The <c>in</c> operator tests if one string is a substring of another:</p>
  <program xml:id="ac_sequenceConditionalOperators_1" label="ac_sequenceConditionalOperators_1" interactive="activecode" language="python">
    <input>
print('p' in 'apple')
print('i' in 'apple')
print('ap' in 'apple')
print('pa' in 'apple')
        </input>
  </program>
  <p>Note that a string is a substring of itself, and the empty string is a
            substring of any other string. (Also note that computer scientists
            like to think about these edge cases quite carefully!)</p>
  <program xml:id="ac_sequenceConditionalOperators_2" label="ac_sequenceConditionalOperators_2" interactive="activecode" language="python">
    <input>
print('a' in 'a')
print('apple' in 'apple')
print('' in 'a')
print('' in 'apple')
        </input>
  </program>
  <p>The <c>not in</c> operator returns the logical opposite result of <c>in</c>.</p>
  <program xml:id="ac_sequenceConditionalOperators_3" label="ac_sequenceConditionalOperators_3" interactive="activecode" language="python">
    <input>
print('x' not in 'apple')
print('a' not in 'apple')
        </input>
  </program>
  <subsection>
    <title>List Membership</title>
  <p><c>in</c> and <c>not in</c> are boolean operators that test membership in a sequence. We
            used them previously with strings and they also work here.</p>
  <program xml:id="ac_sequenceConditionalOperators_4" label="ac_sequenceConditionalOperators_4" interactive="activecode" language="python">
    <input>
fruit = ["apple", "orange", "banana", "cherry"]

print("apple" in fruit)
print("pear" in fruit)
print("watermelon" not in fruit)
        </input>
  </program>

<p> Remember that <c>in</c> and <c>not in</c> only checks for membership in the top level of a list. </p>

  <program xml:id="ac_sequenceConditionalOperators_5" label="ac_sequenceConditionalOperators_5" interactive="activecode" language="python">
    <input>
my_list = ["apple", "orange", "banana", "cherry", ["pear", "kiwi"]]
print("pear" in my_list) # prints False because "pear" is in a sublist
print("g" in my_list) # prints False because "g" is not an member in my_list by itself
print("orange" not in my_list) # prints False because "orange" is a member of my_list 
        </input>
  </program>
  <p>
    <term>Check your understanding</term>
  </p>

  <exercise label="mc_sequenceConditionalOperators_1">
    <statement>
      <p>What is printed by the following statements?</p>
      <program language="python">
        <input>
alist = [3, 67, "cat", [56, 57, "dog"], [ ], 3.14, False]
print(3.14 in alist)
</input>
      </program>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>True</p>
        </statement>
        <feedback>
                    Yes, 3.14 is an item in the list alist.
                </feedback>
      </choice>
      <choice>
        <statement>
          <p>False</p>
        </statement>
        <feedback>
                    There are 7 items in the list, 3.14 is one of them.
                </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="mc_sequenceConditionalOperators_2">
    <statement>
      <p>What is printed by the following statements?</p>
      <program language="python">
        <input>
alist = [3, 67, "cat", [56, 57, "dog"], [ ], 3.14, False]
print(57 in alist)
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>True</p>
        </statement>
        <feedback>
                    in returns True for top level items only.  57 is in a sublist.
                </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>False</p>
        </statement>
        <feedback>
                    Yes, 57 is not a top level item in alist.  It is in a sublist.
                </feedback>
      </choice>
    </choices>
  </exercise>
    <exercise label="ac5_9_2">
    <statement>
      <p>Create one conditional to find whether <q>false</q> is in string <c>str1</c>. If so, assign
        variable <c>output</c> the string <q>False. You aren't you?</q>. Check to see if <q>true</q>
        is in string <c>str1</c> and if it is then assign <q>True! You are you!</q> to the variable <c>
        output</c>. If neither are in <c>str1</c>, assign <q>Neither true nor false!</q> to <c>
        output</c>.</p>
    </statement>
    <program xml:id="ac7_9_2_editor" interactive="activecode" language="python">
      <input>
        str1 = "Today you are you! That is truer than true! There is no one alive who is you-er than you!"

        =====

        from unittest.gui import TestCaseGui

        class myTests(TestCaseGui):

          def testThree(self):
            self.assertEqual(output, "True! You are you!", "Testing that output has the correct value, given the str1 provided.")
            self.assertIn("else", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")
            self.assertIn("elif", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")

        myTests().main()
      </input>
    </program>
  </exercise>
  </subsection>
</section>
