<?xml version="1.0"?>
<exercises xml:id="sequ_exercises">
  <title>Drill and Practice</title>
  <p>
    This section is designed to reinforce your understanding and help you gain confidence in writing
    code.
    Whether you're a beginner or looking to brush up on the basics, these exercises will cover basic
    topics such as variables, data types, statements and expressions.

  </p>
  <p>Take your time with each exercise, and don't hesitate to experiment with different approaches.
    Programming is all about solving problems creatively, so feel free to explore and make
    mistakes—it's all part of the learning process.

  </p>
  <exercise label="ex_seq-1">
   <statement>
    Declare a list called <code>my_lst</code> that includes at least 4 elements (values).
   </statement> 
  
  <program xml:id="ac_seq_1" label="ac_sec_1" interactive="activecode" language="python">
    <input>
       

    </input>
    <tests>
      from unittest.gui import TestCaseGui

      class MyTests(TestCaseGui):
        def testOne(self):
          self.assertTrue('my_lst' in self.getEditorText(), "Testing that you are using the specefied list name.")
          self.assertTrue(len(my_lst) >= 4, "The length of the list should be at least 4")
      
      MyTests().main()
    </tests>
  </program>
  </exercise>

  <exercise label="ex_seq-2">
    <statement>
     Display the length of the list <code>my_stuff</code>.
    </statement> 
   
   <program xml:id="ac_seq_2" label="ac_sec_2" interactive="activecode" language="python">
     <input>
      my_stuff = ["computer", "desk", "chair", "xbox, play station, switch"]
 
     </input>
     <tests>
      from unittest.gui import TestCaseGui

      class MyTests(TestCaseGui):
        def testOne(self):
          self.assertTrue('len(my_stuff)' in self.getEditorText(), "Testing that you are using the correct function and parameter to check the list length.")
          self.assertTrue(self.getOutput().strip() == "4","The expected length of the list is 4")
      
      MyTests().main()

     </tests>
   
   </program>

</exercise>

<exercise label="ex_seq-3">
  <statement>
   Display the last element in the list <code>my_stuff</code>.
  </statement> 
 
 <program xml:id="ac_seq_3" label="ac_sec_3" interactive="activecode" language="python">
   <input>
    my_stuff = ["computer", "desk", "chair", "xbox, play station, switch"]

   </input>
   <tests>
    from unittest.gui import TestCaseGui

    class MyTests(TestCaseGui):
      def testOne(self):
        self.assertTrue('my_stuff[-1]' in self.getEditorText() or 'len(my_stuff)-1' in self.getEditorText() , "Testing that you are referencing the last element correctly.")
        self.assertTrue(self.getOutput().strip() == "xbox, play station, switch","The expected output is the value of the last element in the provided list.")
    
    MyTests().main()

   </tests>
 
 </program>

</exercise>

<exercise label="ex_seq-4">
  <statement>
   Change the value of the first element in the list <c>my_stuff</c> to a value other than "computer".
  </statement> 
 
 <program xml:id="ac_seq_4" label="ac_sec_4" interactive="activecode" language="python">
   <input>
    my_stuff = ["computer", "desk", "chair", "xbox, play station, switch"]

   </input>
   <tests>
    from unittest.gui import TestCaseGui

    class MyTests(TestCaseGui):
      def testOne(self):
        self.assertTrue(my_stuff[0] != 'computer' , "Testing that the value of the first element is not equal to 'computer'.")
    MyTests().main()
   </tests>
 
 </program>

</exercise>
<exercise label="ex_seq-4-1">
  <statement>
Update the fourth element of the list <code>my_stuff</code> so that its value is <c>"game consoles"</c>.
  </statement>
 
 <program xml:id="ac_seq_4-1" label="ac_sec_4-1" interactive="activecode" language="python">
   <input>
    my_stuff = ["computer", "desk", "chair", "xbox, play station, switch", "headphones"]

   </input>
   <tests>
    from unittest.gui import TestCaseGui
    import re
    class MyTests(TestCaseGui):
      def testOne(self):
        self.assertTrue(my_stuff[3] == 'game consoles' , "Testing that the value of the fourth element is equal to 'game consoles'.")
        self.assertTrue(len(my_stuff) == 5, "Testing that the length of my_stuff is 5.")
        self.assertTrue(my_stuff == ["computer", "desk", "chair", "game consoles", "headphones"], "Testing that my_stuff is still a list.")
        self.assertRegex(self.getEditorText(), r"my_stuff\s*\[\s*3\s*\]\s*=\s*['\"]game\s+consoles['\"]", "Testing that the fourth element is updated correctly.")  
    MyTests().main()
   </tests>
 
 </program>
</exercise>

<exercise label="ex_seq-5">
  <statement>
  Complete this program so that it first asks the user to enter a number value between 1 and the length of the list <code>drinks</code>
  then it asks the user to enter a number that represents the desired quantity.  
  Using that input, the program retrieves and displays the corresponding drink information and the total price.
  For example, if the user enters 2 and 3 the output would include "chai latte" and 13.68
  </statement> 
 
 <program xml:id="ac_seq_5" label="ac_sec_5" interactive="activecode" language="python">
   <input>
    drinks = ["tea","chai latte", "latte", "coffee","soda", "berry lemonade"]
    prices = [2.49, 4.56, 3.75, 2.99, 1.99, 3.49]

   </input>
  
  </program>

<solution>
   <program xml:id="ac_seq_5_answer" language="python">
    <input>
      ## question 5 solution ##
      drinks = ["tea","chai latte", "latte", "coffee", "soda", "berry lemonade"]
      prices = [2.49, 4.56, 3.75, 2.99, 1.99, 3.49]

      order = input("Enter a number value between 1 and " + str(len(drinks)) + ": ")
      quantity = input("How many " + drinks[int(order) - 1] + "s would you like? ")

      total = prices[int(order) - 1] * int(quantity)

      print("Your total is","$"+str(total)+".","Your",drinks[int(order) - 1],"will be ready shortly.")
    </input>
  </program>
</solution>
 

</exercise>

  <exercise label="ac8_10_1">
    <statement>
      <p>Write a program that extracts the last three items in the list <c>sports</c> and assigns it to the variable <c>last</c>. Make sure to write your code so that it works no matter how many items are in the list.</p>
    </statement>
    <program xml:id="ac8_10_1_editor" interactive="activecode" language="python">
      <input>
sports = ['cricket', 'football', 'volleyball', 'baseball', 'softball', 'track and field', 'curling', 'ping pong', 'hockey']


=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

  def test_output(self):
     self.assertEqual(last, ['curling', 'ping pong', 'hockey'], "Testing that the value of last is the last three items in sports.")
     self.assertTrue(re.search(r'last\s*=\s*\S*3:', self.getEditorText()), "Hardcode check")


myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac8_10_2">
    <statement>
      <p>Write code that combines the following variables so that the sentence <q>You are doing a great job, keep it up!</q> is assigned to the variable <c>message</c>. Do not edit the values assigned to <c>by</c>, <c>az</c>, <c>io</c>, or <c>qy</c>.</p>
    </statement>
    <program xml:id="ac8_10_2_editor" interactive="activecode" language="python">
      <input>
by = "You are"
az = "doing a great "
io = "job"
qy = "keep it up!"


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(by, 'You are', "Testing original variables.")
    self.assertEqual(az, 'doing a great ', "Testing original variables.")
    self.assertEqual(io, 'job', "Testing original variables.")
    self.assertEqual(qy, 'keep it up!', "Testing original variables.")
    self.assertEqual(message, 'You are doing a great job, keep it up!', "Testing that the value of message is what was expected.")
    self.assertNotIn("You are doing a great job, keep it up!", self.getEditorText(), "Testing for hardcoding (Don't worry about actual and expected values).")


myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="access_ac_2_1_1_11">
    <statement>
      <p>Assign the last element of <c>lst</c> to the variable <c>end_elem</c>. Do this so that it works no matter how long lst is.</p>
    </statement>
    <program xml:id="access_ac_2_1_1_11_editor" interactive="activecode" language="python">
      <input>
lst = ["hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones"]

=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(end_elem, lst[-1], "Testing that end_elem has the correct element assigned.")
      self.assertFalse(re.search(r'end_elem\s*=\s*\S26trombones\S', self.getEditorText()), "Hardcoding Check (Don't worry about actual and expected values)")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="access_ac_2_1_1_12">
    <statement>
      <p>Assign the number of elements in <c>lst</c> to the variable <c>num_lst</c>.</p>
    </statement>
    <program xml:id="access_ac_2_1_1_12_editor" interactive="activecode" language="python">
      <input>
lst = ["hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones"]

=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num_lst, 30, "Testing that num_lst has the correct length assigned.")
      self.assertFalse(re.search(r'num_lst\s*=\s*30', self.getEditorText()), "Hardcoding Check (Don't worry about actual and expected values)")

myTests().main()
        </input>
    </program>
  </exercise>
   <exercise label="question8_10_1">
    <statement>
      <p>What will the output be for the following code?</p>
      <program language="python">
        <input>
let = "z"
let_two = "p"
c = let_two + let
m = c*5
print(m)
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>zpzpzpzpzp</p>
        </statement>
        <feedback>
          <p>The order of concatenation matters.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>zzzzzppppp</p>
        </statement>
        <feedback>
          <p>Think about the order that the program is executed in, what occurs first?</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>pzpzpzpzpz</p>
        </statement>
        <feedback>
          <p>Yes, because let_two was put before let, c has "pz" and then that is repeated five times.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>pppppzzzzz</p>
        </statement>
        <feedback>
          <p>Think about the order that the program is executed in, what occurs first?</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>None of the above, an error will occur.</p>
        </statement>
        <feedback>
          <p>This is correct syntax and no errors will occur.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question8_10_2">
    <statement>
      <p>What will the output be for the following code?</p>
      <program language="python">
        <input>
ls = ['run', 'world', 'travel', 'lights', 'moon', 'baseball', 'sea']
new = ls[2:4]
print(new)
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>['travel', 'lights', 'moon']</p>
        </statement>
        <feedback>
          <p>When we take a slice of something, it includes the item at the first index and excludes the item at the second index.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>['world', 'travel', 'lights']</p>
        </statement>
        <feedback>
          <p>When we take a slice of something, it includes the item at the first index and excludes the item at the second index. Additionally, Python is a zero-index based language.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>['travel', 'lights']</p>
        </statement>
        <feedback>
          <p>Yes, python is a zero-index based language and slices are inclusive of the first index and exclusive of the second.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>['world', 'travel']</p>
        </statement>
        <feedback>
          <p>Python is a zero-index based language.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question8_10_3">
    <statement>
      <p>What is the type of <c>m</c>?</p>
      <program language="python">
        <input>
l = ['w', '7', 0, 9]
m = l[1:2]
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>string</p>
        </statement>
        <feedback>
          <p>Not quite, is it slicing or accessing an element?</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>integer</p>
        </statement>
        <feedback>
          <p>What is happening in the assignment statement for m?</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>float</p>
        </statement>
        <feedback>
          <p>What is happening in the assignment statement for m?</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>list</p>
        </statement>
        <feedback>
          <p>Yes, a slice returns a list no matter how large the slice.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question8_10_4">
    <statement>
      <p>What is the type of <c>m</c>?</p>
      <program language="python">
        <input>
l = ['w', '7', 0, 9]
m = l[1]
</input>
      </program>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>string</p>
        </statement>
        <feedback>
          <p>Yes, the quotes around the number mean that this is a string.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>integer</p>
        </statement>
        <feedback>
          <p>Not quite, look again at what is being extracted.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>float</p>
        </statement>
        <feedback>
          <p>Not quite, look again at what is being extracted.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>list</p>
        </statement>
        <feedback>
          <p>Not quite, is it slicing or accessing an element?</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <query label="practice_feedback_seq">
    <statement>
      <p>Which of the following best describes your experience with these problems?</p>
    </statement>
    <choices>
      <choice>
        <p>I found the problems very confusing and didn't know where to start.</p>
      </choice>
      <choice>
        <p>I started solving the problems but encountered significant difficulties.</p>
      </choice>
      <choice>
        <p>I managed to solve the problems, but it was very challenging.</p>
      </choice>
      <choice>
        <p>I solved the problems with some minor difficulties.</p>
      </choice>
      <choice>
        <statement>
          <p>I found the problems straightforward and solved them easily.</p>
        </statement>
        <feedback>
        </feedback>
      </choice>
    </choices>
  </query>


</exercises>