<?xml version="1.0"?>
<section xml:id="nested-data_chapter-assessment">
  <title>Drill and Practice</title>
  <exercise label="ac17_6_1">
    <statement>
      <p>The variable <c>nested</c> contains a nested list. Assign &#x2018;snake' to the variable <c>output</c> using indexing.</p>
    </statement>
    <program xml:id="ac17_6_1_editor" interactive="activecode" language="python">
      <input>

nested = [['dog', 'cat', 'horse'], ['frog', 'turtle', 'snake', 'gecko'], ['hamster', 'gerbil', 'rat', 'ferret']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(output, "snake", "Testing that output is assigned to correct value")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac17_6_2">
    <statement>
      <p>Below, a list of lists is provided. Use in and not in tests to create variables with Boolean values. See comments for further instructions.</p>
    </statement>
    <program xml:id="ac17_6_2_editor" interactive="activecode" language="python">
      <input>

lst = [['apple', 'orange', 'banana'], [5, 6, 7, 8, 9.9, 10], ['green', 'yellow', 'purple', 'red']]

#Test to see if 'yellow' is in the third list of lst. Save to variable ``yellow``


#Test to see if 4 is in the second list of lst. Save to variable ``four``


#Test to see if 'orange' is in the first element of lst. Save to variable ``orange``

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(yellow, True, "Testing that yellow is assigned to correct value")
   def testTwoB(self):
      self.assertEqual(four, False, "Testing that four is assigned to correct value")
   def testTwoC(self):
      self.assertEqual(orange, True, "Testing that orange is assigned to correct value")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac17_6_3">
    <statement>
      <p>Below, we've provided a list of lists. Use in statements to create variables with Boolean values - see the ActiveCode window for further directions.</p>
    </statement>
    <program xml:id="ac17_6_3_editor" interactive="activecode" language="python">
      <input>

L = [[5, 8, 7], ['hello', 'hi', 'hola'], [6.6, 1.54, 3.99], ['small', 'large']]

# Test if 'hola' is in the list L. Save to variable name test1

# Test if [5, 8, 7] is in the list L. Save to variable name test2

# Test if 6.6 is in the third element of list L. Save to variable name test3

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testA(self):
      self.assertEqual(test1, False, "Testing that test1 has the correct value.")
   def testB(self):
      self.assertEqual(test2, True, "Testing that test2 has the correct value.")
   def testC(self):
      self.assertEqual(test3, True, "Testing that test3 has the correct value.")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac17_6_8">
    <statement>
      <p>Iterate through the contents of <c>l_of_l</c> and assign the third element of sublist to a new list called <c>third</c>.</p>
    </statement>
    <program xml:id="ac17_6_8_editor" interactive="activecode" language="python">
      <input>

l_of_l = [['purple', 'mauve', 'blue'], ['red', 'maroon', 'blood orange', 'crimson'], ['sea green', 'cornflower', 'lavender', 'indigo'], ['yellow', 'amarillo', 'mac n cheese', 'golden rod']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(third, ['blue', 'blood orange', 'lavender', 'mac n cheese'], "Testing that third has the correct list assigned to it.")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac17_6_9">
    <statement>
      <p>Given below is a list of lists of athletes. Create a list, <c>t</c>, that saves only the athlete's name if it contains the letter <q>t</q>. If it does not contain the letter <q>t</q>, save the athlete name into list <c>other</c>.</p>
    </statement>
    <program xml:id="ac17_6_9_editor" interactive="activecode" language="python">
      <input>

athletes = [['Phelps', 'Lochte', 'Schooling', 'Ledecky', 'Franklin'], ['Felix', 'Bolt', 'Gardner', 'Eaton'], ['Biles', 'Douglas', 'Hamm', 'Raisman', 'Mikulak', 'Dalton']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(t, ['Lochte', 'Bolt', 'Eaton', 'Dalton'], "Testing that t is assigned to correct values.")
   def testFiveA(self):
      self.assertEqual(other, ['Phelps', 'Schooling', 'Ledecky', 'Franklin', 'Felix', 'Gardner', 'Biles', 'Douglas', 'Hamm', 'Raisman', 'Mikulak'], "Testing that other is assigned to correct values.")

myTests().main()
        </input>
    </program>
  </exercise>
</section>
